# DS4DS Homework Exercise Sheet 01

We have the two matrices $A$, $B$:

$
A = \begin{bmatrix} 0 & 0 \\ 2 & 0 \end{bmatrix}, \\
B = \begin{bmatrix} 0 & -1 \\ 4 & 0 \end{bmatrix}
$

```julia
A = [0 0; 2 0]

B = [0 -1; 4 0]
```

### Task 1: Calculate the eigenvalues of the matrices $A$,$B$ and assign them to *eig_A* and *eig_B* - (1 point)
**Hint:** You should calculate them by hand and insert them here. Alternatively, you may solve this in Julia, but without the use of any extra packages.

<!-- grade_id: cell-a42b9d1c810ef1f9, points: 0, grade: false, solution: true -->
```julia
eig_A = 0
eig_B = 0

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-75f64657aa33a44d, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(eig_A,Array)
@assert isa(eig_A,Vector)
@assert !isa(eig_A,Matrix)
@assert length(eig_A) == 2

```

<!-- grade_id: cell-a100be41157a8b07, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(eig_B,Array)
@assert isa(eig_B,Vector)
@assert !isa(eig_B,Matrix)
@assert length(eig_B) == 2

```

### Task 2: Calculate the matrix exponential function of matrix A and B (by hand) [ exp(A\*t) and exp(B\*t) ] - (3 points)

**Hint 1:** You should calculate them by hand and insert them here. Alternatively, you may solve this in Julia, but without the use of any extra packages.

**Hint 2:** The function should at the end return a 2x2 matrix 

**Important: The use of exp() function will result in 0 points**

```julia
#Here is an example how the functions should look like
function example(t)
    return [1*t 2; 3 4*t]
end
```

<!-- grade_id: cell-86de8e669e0bc430, points: 0, grade: false, solution: true -->
```julia
function matrixExponentA(t)
    return nothing
end

function matrixExponentB(t)
    return nothing
end

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-1b2cc81a9c556ff4, points: 1.5, grade: true, solution: false -->
```julia
@assert isa(matrixExponentA,Function)
@assert size(matrixExponentA(0)) == (2,2)

```

<!-- grade_id: cell-144af36570832246, points: 1.5, grade: true, solution: false -->
```julia
@assert isa(matrixExponentB,Function)
@assert size(matrixExponentB(0)) == (2,2)

```

<!-- grade_id: cell-9908aec96171e3a5, points: 0, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

### Task 3: Calculate the system reactions $x(t)$ for a general $t$ and $x_0$ - (2 points)

```julia
x0 = [1; 1]

#Here is an example how the functions should look like. The function should at the end return a 2x1 vector
function example(t, x0)
    return [1 ; 2]
end
```

<!-- grade_id: cell-cb80f70d0d95495c, points: 0, grade: false, solution: true -->
```julia
function state_response_A(t, x0)
    return nothing
end

function state_response_B(t, x0)
    return nothing
end

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-37d579c2696587c5, points: 1, grade: true, solution: false -->
```julia
@assert isa(state_response_A,Function)
@assert size(state_response_A(0, [0; 0])) == (2,)

```

<!-- grade_id: cell-915b2d57087b0936, points: 1, grade: true, solution: false -->
```julia
@assert isa(state_response_B,Function)
@assert size(state_response_B(0, [0; 0])) == (2,)

```

### Task 4: Plot the resulting time series with plots.jl for a given time interval into a common plot - (3 points)

Note: Take care of proper formatting of the plot. In particular, add correct axis labeling and a legend to be able to differentiate the time series.

```julia
using Plots
```

```julia
# time points
t_values = 0:0.01:5
N = size(t_values)[1]

x0 = [1; 1]

# fill with time series for matrix A and B respectively
xt_A = zeros((N, 2));
xt_B = zeros((N, 2));
```

<!-- grade_id: cell-7059f3b5672fbebf, points: 0, grade: false, solution: true -->
```julia
#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

```julia
@assert size(xt_A) == (501, 2)
@assert size(xt_B) == (501, 2)

@assert any(xt_A .!= 0)
@assert any(xt_B .!= 0)
```

<!-- grade_id: cell-3cda36de24666aa7, points: 1.5, grade: true, solution: true -->
```julia
#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-7dd6c12dee466781, points: 1.5, grade: true, solution: true -->
```julia
#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

### Task 5: In each case, make a statement whether the autonomous system described by the matrices is state-stable - (1 point)
Therefore assign the corresponding value to the given variables.

The system is:\
1: Stable\
2: asymptotically stable\
3: Unstable

<!-- grade_id: cell-a8807b7c513764fa, points: 0, grade: false, solution: true -->
```julia
stability_of_A = 0
stability_of_B = 0

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-b535e209a9bf3054, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(stability_of_A, Number)

```

<!-- grade_id: cell-d23001987877a646, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(stability_of_B, Number)

```