# DS4DS Homework Exercise Sheet 04

### Task 1: Least Squares Linear Approximation - (2 points)

The measurements shown in the figure below were gathered from an unknown process.

<img src="measurements_updated.png" alt="measurements" width="400"/>

However, we can assume a linear relationship between the measurements $y$ and the inputs $u$ for the underlying process

$$
\begin{align}
    y[k] = w_0 + w_1 u[k]
\end{align}.
$$

Your task is to identify the parameter vector $$\boldsymbol{w} = \begin{bmatrix} w_0\\ w_1\end{bmatrix} $$ using the least squares approach shown in the lecture.

**Note:** You can solve this problem using either pen and paper or Julia. Place your calculated values for $w_0$ and $w_1$ into the variables below.

**Tip:** Think about the role of the two parameters you have acquired. Do they make sense for the data? Is there a way to visualize your result?

<!-- grade_id: cell-2f711984831d75f9, points: 0, grade: false, solution: true -->
```julia
w_0 = NaN
w_1 = NaN

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-b24199f2e9388d9c, points: 1, grade: true, solution: false -->
```julia
@assert isnan(w_0) == false "Your value for w_0 is NaN"

# please leave this cell as it is

```

<!-- grade_id: cell-fee1dd5999288431, points: 1, grade: true, solution: false -->
```julia
@assert isnan(w_1) == false "Your value for w_1 is NaN"

# please leave this cell as it is

```

### Task 2: Parameter Estimation of the DC Motor - (5 points)

The electrical behaviour of the DC motor in steady state is characterized by the following static equation

$$
\begin{align}
    i_s = u_s \frac{1}{R_s} - \omega \frac{\Psi_E}{R_s}
\end{align}
$$

where $u_s$ is the voltage provided by a power supply, $\omega$ is the rotation speed which is set by a mechanical load machine, $R_s$ and $\Psi_E$ are unknown values and $i_s$ is the current that can be measured after the motor reached steady-state operation. For the whole task, you can assume that $u_s$ and $\omega$ are perfectly known, while the measured values for $i_s$ are affected by bias-free, additive, Gaussian noise.

**a)** Consider the system described above. Assign each of the equation components to the different elements of the least squares problem. For this, map each of the components in the cell below to one of the following classes: **- (0.5 points)**
- class **1**: Measurements $\boldsymbol{y}$
- class **2**: Regressors $\boldsymbol{z}$
- class **3**: Parameters $\boldsymbol{w}$ (this is the letter $\boldsymbol{w}$ for the parameter vector not to be confused with the greek symbol $\omega$ for the angular velocity)

<!-- grade_id: cell-6d3f3c8637b1dd1f, points: 0, grade: false, solution: true -->
```julia
# choose 1,2 or 3

class_i_s = 0
class_u_s = 0
class_one_over_R_s = 0  # i.e. 1/R_s
class_omega = 0
class_minus_Psi_E_over_R_s = 0  # i.e. - Psi_E/R_s

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-18008ded3921637c, points: 0.5, grade: true, solution: false -->
```julia
@assert 1 <= class_i_s <= 3  
@assert 1 <= class_u_s <= 3  
@assert 1 <= class_one_over_Rs <= 3  
@assert 1 <= class_omega <= 3  
@assert 1 <= class_minus_Psi_E_over_R_s <= 3  

# please leave this cell as it is

```

**b)** Implement the function below to calculate the parameters $\boldsymbol{w}$ from the regressor matrix $\boldsymbol{Z}$ and the measurement vector $\boldsymbol{y}$ using the ordinary least squares approach from the lecture. **- (1 point)**

**Hint**: This is not specific to the system at hand.

<!-- grade_id: cell-7a395d03c399fb11, points: 0, grade: false, solution: true -->
```julia
function parameter_calculation_OLS(Z, y)
    """
    Args:
        Z: regressor matrix with shape (n_measurements, n_regressors)
        y: measurement_vector with shape (n_measurements)
    
    Returns:
        The parameter vector w with shape (n_parameters)
    """
    
    #--- YOUR CODE STARTS HERE ---#
    
    #--- YOUR CODE ENDS HERE ---#
    return w
end;
```

<!-- grade_id: cell-afe2c2a7f1fc95be, points: 1, grade: true, solution: false -->
```julia
@assert isa(parameter_calculation_OLS, Function)

# please leave this cell as it is

```

<!-- grade_id: cell-23d6592d2175e406, points: 0, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

**c)** Use the function implemented in subtask b) and the given data for $u_s$, $\omega$ and $i_s$ to estimate values for $a = \frac{1}{R_s}$ and $b = -\frac{\Psi_E}{R_s}$. Put your results for $a$ and $b$ into the variables below. **- (1 point)**

```julia
using MAT

data_task_2 = matopen("data_task_2.mat")

# read data
u_s = read(data_task_2, "u_s");
omega = read(data_task_2, "omega");
i_s = read(data_task_2, "i");
```

<!-- grade_id: cell-fb932c68d8a13b80, points: 0, grade: false, solution: true -->
```julia
# variables for you to overwrite
a = NaN
b = NaN

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-eec824a224d98d41, points: 1, grade: true, solution: false -->
```julia
@assert isnan(a) == false "Your value for a is NaN"
@assert isnan(b) == false "Your value for b is NaN"

@assert a > 0 "The resistance must be positive by physical constraints, that is, 'a' must be positive"
@assert b < 0 "The resistance and the flux linkage must be positive, that is, 'b' must be negative"

# please leave this cell as it is

```

<!-- grade_id: cell-3b7f861ee03b4c2b, points: 0, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

**d)** Implement a function to compute estimations for $R_s$ and $\Psi_E$ based on $a$ and $b$. **- (0.5 points)**

<!-- grade_id: cell-de734069667417d5, points: 0, grade: false, solution: true -->
```julia
function compute_Rs_Psi_E(a, b)

    #--- YOUR CODE STARTS HERE ---#
    
    #--- YOUR CODE ENDS HERE ---#

    return R_s, Psi_E
end;
```

<!-- grade_id: cell-73c4a9a074fd3837, points: 0.5, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

**e)** Analyse the accuracy of your parameter estimates. For this, take into consideration that the measured values for $i_{s}$, that you loaded from the file, are affected by bias-free, additive, Gaussian noise:

$$
\begin{align}
i_{s} = i_{s, true} + \nu \quad \mathrm{with} \, \nu \sim \mathcal{N}(0, \sigma_n^2).
\end{align}
$$

i) Estimate the corrected sample variance of $\nu$ from the measurements for $i_s$, that you were given, and your estimate $\hat{i}_s$, that is based on the estimated parameters. **- (0.5 points)**

<!-- grade_id: cell-58f9cd149f44ba50, points: 0, grade: false, solution: true -->
```julia
var_nu = NaN

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-005a85dce295b7a4, points: 0.5, grade: true, solution: false -->
```julia
@assert isnan(var_nu) == false "Your value for the variance of nu is not a number."
@assert var_nu > 0 "A variance is always positive."

# please leave this cell as it is

```

ii) Implement the function below to estimate the covariance matrix of the parameter vector $\boldsymbol{w}$ based on the noise variance $\sigma_n^2$ and the regressor matrix $\boldsymbol{Z}$ **- (0.5 points)**

<!-- grade_id: cell-1123100767a66d44, points: 0, grade: false, solution: true -->
```julia
function estimate_covariance_matrix(var_nu, Z)
    """
    Args:
        var_nu: variance of the noise process
        Z: regressor matrix with shape (n_measurements, n_regressors)

    Returns:
        covariance matrix of w
    """
    
    #--- YOUR CODE STARTS HERE ---#
    
    #--- YOUR CODE ENDS HERE ---#

    return cov_w
end;
```

<!-- grade_id: cell-ebe32496af8bba27, points: 0.5, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

iii) Evaluate the coefficient of determination for the regression problem (see https://en.wikipedia.org/wiki/Coefficient_of_determination). **- (1 point)**

$$
\begin{align}
    R^2 = 1 - \frac{\sum_i{(y_i - \hat{y}_i)^2}}{\sum_i{(y_i - \bar{y})^2}}
\end{align}
$$

<!-- grade_id: cell-b0cecce9d9deca39, points: 0, grade: false, solution: true -->
```julia
R_squared = NaN  # put your solution into this variable

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-5fe000009068f83d, points: 1, grade: true, solution: false -->
```julia
@assert isnan(R_squared) == false "Your value for the coefficient of determination is not a number."

# please leave this cell as it is

```

### Task 3: Estimate Polynom Coefficients - (3 points)

The given data was produced using a polynomial of the form


$$
\begin{align}
    y = w_0 + x \cdot w_1 + \dots + x^p \cdot w_p
\end{align}
$$

with an unknown degree $p$.

Below you are given data for the inputs $x$ and the outputs $y$. On the one hand you are given training data you will use for the identification of coefficients (estimation of a model) and validation data
that is used to evaluate the generalization performance of your model. All of the data for the measurements $y$ is affected by bias-free, additive, Gaussian noise, i.e., the values you are given below do not fully follow the formula given above, but instead

$$
\begin{align}
    y = w_0 + x \cdot w_1 + \dots + x^p \cdot w_p + \nu \quad \mathrm{with} \, \nu \sim \mathcal{N}(0, \sigma_n^2).
\end{align}
$$

The variance of the noise process is unknown, but it is also not necessary to estimate it for completing this task.

```julia
using MAT
data_task_3 = matopen("data_task_3.mat");
```

```julia
# training data
x_training = read(data_task_3, "x_training")[1, :];
y_training = read(data_task_3, "y_training")[1, :];
 
# validation data
x_validation = read(data_task_3, "x_validation")[1, :];
y_validation = read(data_task_3, "y_validation")[1, :];
```

**a)** Consider the regression problem of the form:

$$
\begin{align}
    \boldsymbol{y} = \boldsymbol{Z} \cdot \boldsymbol{w}  + \boldsymbol{\nu}.
\end{align}
$$

Implement the function below to compute the regressor matrix $\boldsymbol{Z}$ for the polynom coefficient estimation problem described above. The function takes the input data $\boldsymbol{x}$ and the degree of the polynom $p$ as inputs. **- (0.5 points)**

**Example:** If $p=2$, we get a model until the 2nd polynomial degree

$$
\begin{align}
    y = w_0 + x \cdot w_1 + x^2 \cdot w_2
\end{align}
$$

which results in the regressor vector

$$
\begin{align}
    \boldsymbol{z} = \begin{bmatrix} 1 & x & x^2 \end{bmatrix}.
\end{align}
$$

With an input vector $\boldsymbol{x} = \begin{bmatrix} 1 & 2 & 3 \end{bmatrix}$ this results in the regressor matrix

$$
\begin{align}
    \boldsymbol{Z} = \begin{bmatrix}
        1 & 1 & 1 \\
        1 & 2 & 4 \\
        1 & 3 & 9 \\
    \end{bmatrix}.
\end{align}
$$



<!-- grade_id: cell-ef3d22676d057b8e, points: 0, grade: false, solution: true -->
```julia
function regressor_matrix(x, p)
    """
    Args:
        x: Input data
        p: Degree of the polynomial
    
    Returns:
        Regressor matrix Z with shape (length(x), p+1)
    """
    
    #--- YOUR CODE STARTS HERE ---#
    
    #--- YOUR CODE ENDS HERE ---#

    return Z
end
```

<!-- grade_id: cell-4075f6537a5a6334, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(regressor_matrix, Function)

# please leave this cell as it is

```

<!-- grade_id: cell-a7e659c0a88fbb76, points: 0, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

**b)** Write the computation of the parameter estimation using OLS as a function of the input data $\boldsymbol{x}$, the output data $\boldsymbol{y}$ and the degree of the polynomial $p$. Your goal is to find the parameters $\boldsymbol{w}$ for the polynom of degree $p$. Use the function for the computation of the regression matrix that was asked for in subtask a) **- (0.5 points)**

<!-- grade_id: cell-be54ff5be59de6b8, points: 0, grade: false, solution: true -->
```julia
function compute_parameter_estimation(x, y, p)
    """
    Args:
        x: Input data
        y: Output data
        p: Degree of the polynomial
    
    Returns:
        parameter vector w with shape (p+1)
    """
    
    #--- YOUR CODE STARTS HERE ---#
    
    #--- YOUR CODE ENDS HERE ---#

    return w
end
```

<!-- grade_id: cell-1446419df57805ac, points: 0.5, grade: true, solution: false -->
```julia
@assert isa(compute_parameter_estimation, Function)

# please leave this cell as it is

```

<!-- grade_id: cell-36a02ed893a4bb07, points: 0, grade: true, solution: false -->
```julia
# please leave this cell as it is

```

**c)** Your goal now is to find the fit onto the training data that best generalizes to the validation data. **- (2 points)**

For this:
- iterate over the degree of the polynomial $p$,
- find the parameter vector $\boldsymbol{w}_p$ based on the training data $\boldsymbol{x}_{training}$ and $\boldsymbol{y}_{training}$,
- compute the estimates $\hat{\boldsymbol{y}}$ using $\boldsymbol{w}_p$ and $\boldsymbol{x}_{validation}$, and
- compare the estimates $\hat{\boldsymbol{y}}$ with $\boldsymbol{y}_{validation}$ using the given error function.

```julia
function error_function(a, b)
    """Computes the error between the two input vectors."""

    @assert length(a) == length(b)
    @assert ndims(a) == ndims(b) == 1
    
    N = length(a)
    return sqrt(1/N * sum((a-b).^2))
end
```

Find a polynomial model with **an error below $10$** according to the steps described above and report its degree in the variable $p_{sol}$ below. 

**Hint 1:** $p_{sol} \in [0, 10]$

**Hint 2:** Multiple degrees for the polynomial produce an error that is low enough. 

<!-- grade_id: cell-de7a80003acab8b5, points: 0, grade: false, solution: true -->
```julia
p_sol = NaN

#--- YOUR CODE STARTS HERE ---#

#--- YOUR CODE ENDS HERE ---#
```

<!-- grade_id: cell-fefbefb113ecb7c9, points: 2, grade: true, solution: false -->
```julia
@assert isnan(p_sol) == false "Your value for the p_final is not a number."
@assert 0 <= p_sol <= 10

# please leave this cell as it is

```

**For the interested (no extra points or credit):** Plot the estimates of the model, the training data and the validation data for each value of $p \in [0, 10]$. Why is the highest degree of freedom not necessarily the best choice? What effect that you might know from the context of machine learning can be seen here?

Please remove your plotting commands from the notebook before submission. Thank you!

```julia

```